2. Algorithmus zur Schaltjahrbestimmung (6 Punkte)
In unserem Kalender sind in regelmäßigen Abständen Schaltjahre eingefügt. Zur exakten Feststellung, ob
ein Jahr ein Schaltjahr ist, dienen folgende Regeln:

a) Ist die Jahreszahl durch 4 teilbar, so ist das Jahr ein Schaltjahr.
b) Regel a) hat allerdings die Ausnahme, dass ein Jahr das durch 100 teilbar ist, kein Schaltjahr ist.
c) Regel b) hat allerdings die Ausnahme, dass ein Jahr das durch 400 teilbar ist, doch ein Schaltjahr ist.

Entwickeln Sie dafür einen Algorithmus.
Verwenden Sie dafür folgedne Signatur: isLeapYear(↓year, ↑isLeapYear)

<!-- grade_id: cell-0b23f7c8abd6c6ea1, points: 0, grade: false, solution: true -->
```C++11
bool isLeapYear (int year) {
    bool isLeapYear = false;

    // BEGIN SOLUTION
    if (year % 4 == 0) {
        isLeapYear = true;
        if (year % 100 == 0) {
            isLeapYear = false;
            if (year % 400 == 0) {
                isLeapYear = true;
            }
        }
    }
    // END SOLUTION
    return isLeapYear;
}
```

<!-- grade_id: cell-c92cd296fc59db5f, points: 6, grade: true, solution: false -->
```C++11
assert(NULL);
assert(isLeapYear(2000) == 1);
assert(isLeapYear(2019) == 0);
assert(isLeapYear(2020) == 1);

int year = 2000;
printf("%d %s", year, isLeapYear(year) ? "is a leapyear" : "is no leap year");
```

Output:
```
2000 is a leapyear
```